# Part A

This cell is meant to be read-only. Students won't be able to edit this.

Answer the following questions?
What is your name and net id? 

<!-- grade_id: Manual_graded_cell, points: 1, grade: true, solution: true -->


Fix the following code to return a list of odd numbers until a given number n. Type cast the given number to int, if it is a float. Don't miss to handle corner cases.

<!-- grade_id: Auto_graded_cell, points: 0, grade: false, solution: true -->
```python
def odd_number_generator(n):
    odd_numbers_list = []
    ### BEGIN SOLUTION
    if type(n) not in [int, float]:
        return odd_numbers_list
    else:
        for number in range(1, int(n)+1):
            if number%2 == 1:
                odd_numbers_list.append(number)
    ### END SOLUTION
    return odd_numbers_list
```

<!-- grade_id: Assert_Statements, points: 3, grade: true, solution: false -->
```python
assert odd_number_generator(5) == [1, 3, 5]
assert odd_number_generator("I") == []
assert odd_number_generator(False) == []
```

Write a pseudo-code to reverse a given number n:
Make sure to handle corner cases